# __Programming Assignment 5__

## __Assignment Overview__

Write a program that takes user input of 5 integers, stores them in a list and then loops through the list to calculate the average of those 5 numbers as the output. Additional intructions are added in the comments below, make sure to read through them carefully.

Be sure to test your code against the test cases! Submit your code once your tests are successfully passing.

### Lets start coding!! 

#### Using this Notebook

* You should begin at the top and execute each code cell, adding your code when you see # YOUR CODE HERE.
* Notice that some of the cells are not able to be edited – they are just meant to be run to test your code, and perform the auto grading. 
* When you execute the TEST CASES cell, what you want to see is no output. You may want to intentionally make the function give a wrong result, so that you can see what the error message looks like – and then fix your function. Once you execute the TEST CASES cell, and get no output, then your function is most likely correct!

#### First task: Prompt the user to enter numbers and store their input in a list

First, you will need to ask the user to enter 5 numbers and store them in a list. [**Hint** : You will probably want to use a *for* loop here] 

<!-- grade_id: cell-7a82fef6d1210fe2, points: 0, grade: false, solution: true -->
```python
### GRADED SECTION 1 : prompt the user to enter 5 numbers and store them 
### in a list called homework_grades 

# REMEMBER to remove the `raise NotImplementedError()` statement when you
# have added your code.  If you do not remove this statement, your code
# will fail the autograding tests.
# (This line is included here to ensure the functioning of the autograder.)

homework_grades = []

for count in range(0, 5):
    # get grade from the user
    grade = int(input('Please enter a grade: '))
    
    # add grade to the list
    homework_grades.append(grade)
```

Output:
```
Please enter a grade:  100
Please enter a grade:  85
Please enter a grade:  30
Please enter a grade:  84
Please enter a grade:  56
```

```python
#print out the list of numbers 

print("You entered the following homework grades: ", homework_grades)
```

Output:
```
You entered the following homework grades:  [100, 85, 30, 84, 56]
```

#### Second Task: Calculating the average 

Now, use this list to calculate the average of all numbers entered.  In this case, you are **required to use a for loop** when doing the calculation (this is one of the learning objectives for this assignment!).

If you look at the code below, a function has been outlined for you, with two areas for you to add code to calculate the average.  First, you will add up all the numbers in the list.  Then you will calculate the average by dividing by 5 (how many numbers are in the list).

**Hint:** It is very common for beginning programmers to want to use print() instead of return.  Here, we _don't_ want this function to actually print the result -- we just want it to **return** the right answer, so that another section of code could print it, if we wanted to.

Write your code in the following cell, execute the cell, and then scroll down to the unit tests that the autograder will use -- and execute them to see if your solution to this function is correct.

<!-- grade_id: cell-3bd0d4a350efc12f, points: 0, grade: false, solution: true -->
```python
# Use a for loop (with range()) to add up all of the grades in the list. 
# We have created a function that takes the list of numbers as parameter 
# and contains a variable named average already. Add your code in the mentioned spaces

def calc_average(grade_list) :

    average = 0
    
    # REMEMBER to remove the `raise NotImplementedError()` statement when you
    # have added your code.  If you do not remove this statement, your code
    # will fail the autograding tests.
    # (This line is included here to ensure the functioning of the autograder.)
    total = 0
    
    for index in range(0, 5):
        total = total + grade_list[index]
    
    ## After the for loop finishes, divide the total by 5, in order to calculate the average.
    
    average = total / 5 
    
    return average

```

### Testing your work 

Now, execute each cell below.  For the first cell, you can change the values stored in the variable *your_list* and then think about what the correct answer would be.

The following two cells that are labeled "GRADED TEST CASES". You should execute these cell.  The first test case verifies that your function is calculating the average correctly.  The second test case verifies that you used a *for* loop, as is required. If you see no output, then your code passed the tests!

One you are passing all tests here, **submit** your work to be graded using the "Submit Assignment" button.

```python
# UNGRADED TEST CASES : You can use your own test lists to validate your code 
your_list = [100,85,30,84,56]
calc_average(your_list)
```

Output:
```
71.0
```

<!-- grade_id: cell-49764892bbc6c937, points: 5, grade: true, solution: false -->
```python
# GRADED TEST CASES 1: You will not be able to edit this cell, but can run it for testing
'''Test your solution for autograding with the tests added below in this cell'''
### BEGIN TESTS
assert calc_average([1,3,4,5,6]) == 3.8 , "Please check your code for calculating the average ." 
assert calc_average([0,4,7,19,5]) == 7 , "Please check your code for calculating the average ."
### END TESTS
```

<!-- grade_id: cell-456f37f799c7c183, points: 5, grade: true, solution: false -->
```python
# GRADED TEST CASES 2: You will not be able to edit this cell, but can run it for testing

import inspect

# Check - for loop implementation
for_loop_found = False

calc_avg_lines = inspect.getsource(calc_average)
code_lines = calc_avg_lines.splitlines()

for line in code_lines:
    if "for" in line and "in" in line and ":" in line:
        for_loop_found = True
        
### BEGIN TESTS
assert for_loop_found == True, "Please verify you've implemented a for loop in calc_average."
### END TESTS
```

### Final Step :  Submit your Work for Grading

Great job on your work this week! To receive a final score for your work, please select the "Submit Assignment" button at the top of your lab.
